.. _Final_Exam_W15/q_17:

Final Exam Fall 2017 - Question 17
----------------------------------


.. activecode:: Final_Exam_W15_q_17
   :language: python
   :autograde: unittest
   :chatcodes:
   :practice: T
   :topics: Final_Exam_W15/q_17

   ``instructors_d`` has information on instructors at U of M, including the number of instructors and information on each
   instructor. For each instructor, there is information on the instructor's name, uniqname, and department(s) they
   teach in. Only information for the first three instructors are shown.

   Write one line of code that uses ``instructors_d`` to store the (3-item) list of instructor dictionaries into the
   variable ``instructors_lst``.

   Afterwards, write code to create a list of uniqnames for instructors that teach in SI. (If they teach in SI,
   one of the strings in the instructor’s ``dept`` list is the string ``"SI"``). Save this list into the variable ``SI_lst``,
   which should be: ``['soney', 'alexturn']``.

   ~~~~
   instructors_d = {
    "university": "University of Michigan",
    "num_instructors": 3,
    "instructors": [
    {
    "first_name": "Steve",
    "last_name": "Oney",
    "uniqname": "soney",
    "dept": ["SI"]
    },
    {
    "first_name": "Claudia",
    "last_name": "Shannon",
    "uniqname": "shclaude",
    "dept": ["ASIAN"]
    },
    {
    "first_name": "Alexandra",
    "last_name": "Turning",
    "uniqname": "alexturn",
    "dept": ["EECS", "SI"]
    }
    ]
   }

   # Write your code here

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

     def testOne(self):
         self.assertEqual(instructors_lst, instructors_d['instructors'], "Testing whether you've got the correct value in instructors_lst.")
         self.assertEqual(SI_lst, ['soney', 'alexturn'], "Testing whether you've got the correct value in SI_lst.")


   myTests().main()
