.. _Final_Exam_W15/q_22:

Final Exam Fall 2017 - Question 22
----------------------------------


.. activecode:: Final_Exam_W15_q_22
   :language: python
   :autograde: unittest
   :chatcodes:
   :practice: T
   :topics: Final_Exam_W15/q_22

   Write code that sorts the dictionary below in ascending order by value and assign the sorted **keys** to the
   variable ``d_sorted``.

   ~~~~
   d = {'Scandal': 7, 'Lost': 6, 'Riverdale': 1, 'Game of Thrones': 8, 'Stranger Things': 2}

   sortedShows =

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

     def testOne(self):
         self.assertEqual(sortedShows, ["Riverdale", "Stranger Things", "Lost", "Scandal", "Game of Thrones"],
         "Testing the value of d_sorted.")

   myTests().main()
