.. _Final_Exam_F15/q_06:

Final Exam Fall 2015 - Question 06
------------------------------------


.. activecode:: Final_Exam_F15_q_06
   :language: python
   :autograde: unittest
   :chatcodes:
   :practice: T
   :topics: Final_Exam_F15/q_06

   Write a function **popular_tweeters** that will take a variable like s and return
   a list of the screen_names in order based on their followers_count.

   ~~~~
   s = [
      {
         "id": 672149962780749800,
         "user": {
            "followers_count": 823,
            "screen_name": "alimahrer",
            "id": 156844181
         }
      },
      {
         "id": 672149247811313700,
         "user": {
            "followers_count": 1569,
            "screen_name": "CatchaJob_us",
            "id": 2348426863
         }
      },
      {
         "id": 672148783413731300,
         "user": {
            "followers_count": 15575,
            "screen_name": "GymCastic",
            "id": 812521788
         }
      }
   ]



   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

     def testOne(self):
         self.assertEqual(popular_tweeters(s), ['GymCastic', 'CatchaJob_us', 'alimahrer'],
                          "Testing the value of popular_tweeters(s).")

   myTests().main()
