.. _Final_Exam_F17/q_14:

Final Exam Winter 2015 - Question 14
------------------------------------


.. activecode:: Final_Exam_F17_q_14
   :language: python
   :autograde: unittest
   :chatcodes:
   :practice: T
   :topics: Final_Exam_F17/q_14

   Define a function f that takes a list of strings and returns a list containing the first letter of every word that contains the letter z. Make it pass the test below. Your function **must use** manual accumulation (i.e., you may not use map, filter, reduce, zip, or list comprehensions.)

   ~~~~



   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

     def testOne(self):
         correct_results = ['w','i','n']
         results = f(['whiz', 'success', 'in the zone', 'yep', 'not another pizza!'])
         self.assertEqual(results, correct_results, "Testing whether `f(['whiz', 'success', 'in the zone', 'yep', 'not another pizza!'])` returns the correct value.")

   myTests().main()
